# Module 6 Assignment


A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)


```python
% matplotlib inline

import numpy as np
import scipy.stats as stats

import pandas as pd
import seaborn as sns

from nose.tools import assert_almost_equal, assert_equal, assert_is_instance, assert_is_not

import matplotlib as mpl
import matplotlib.pyplot as plt
```

# Problem 1: Calculating a Simple Probability

Write a function called `probability_calc` that takes in the number of observed heads from a coin flip experiment and the total number of coin flips, and returns the probability that the coin lands on heads.

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
def probability_calc(num_heads,n):
    """
    Inputs
    ------
    
    num_heads: an integer, the number of heads observed
    
    n: an integer, the total number of coin flips
    
    Output
    ------
    
    prob: the probability of the coin landing on heads
    """    
    
    ### YOUR CODE HERE
    
    return prob
```

<!-- grade_id: p1_test, points: 1, grade: true, solution: false -->
```python
my_prob = probability_calc(37, 233)
assert_almost_equal(my_prob, 0.15879828326180256)
```

# Problem 2: Using Bayes Theorem

Write a function called `bayes_calc` which takes in the probability of event $A$, and the probability of event $B$ and return $P(A|B)$. Assume that $A$ and $B$ are independent events.

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
def bayes_calc(prob_A, prob_B):
    """
    Inputs
    -------
    
    prob_A: a float, the probability of event A
    
    prob_B: a float, the probability of event B
    
    Output
    ------
    
    prob: the probability of event A given event B
    
    """
    
    ### YOUR CODE HERE
    
    return prob
```

<!-- grade_id: p2_test, points: 1, grade: true, solution: false -->
```python
my_prob = bayes_calc(15 / 66, 13 / 68)
assert_almost_equal(my_prob, 0.227272, places=5)
```

```python
# Read in dow jones index dataset
df = pd.read_csv('./dow_jones_index.data')
df.head()
```

# Problem 3: Possion distribution of columns from dow jones dataset
In this problem you'll finish writing the `plot_poisson` function. A dataframe of the dow jones dataset and the column name are passed into this function. Your task is to:
- plot the PMF of the Poisson distribution for column `col` using the `bar` function from the Axes object.
    - Use the mean of the column to create the Poisson distirbution
    - The number of points is given by `n_pts`
- Your plot should contain a label for the x and y axis and also have a title.
- Return the `axes` object

<!-- grade_id: p4_answer, points: 0, grade: false, solution: true -->
```python
def plot_poisson(df, col='open', n_pts=100):
    '''
    df: dataframe containg data from dow jones index
    col: column name in df variable
    n_pts: number of points
    returns axes object
    '''
    # Define plot layout
    fig, ax = plt.subplots(figsize=(10, 5))
    
    ### YOUR CODE HERE
    return ax
```

<!-- grade_id: p4_run, points: 1, grade: true, solution: false -->
```python
ch = plot_poisson(df,'open')
```

<!-- grade_id: p4_test, points: 1, grade: true, solution: false -->
```python
from helper import pp
print("Your Plot should look similar to this:")
sol = pp(df)

assert_is_instance(ch, mpl.axes.Axes, msg='Return a Axes object.')  
assert_is_instance(ch, mpl.axes.Axes, msg='Return a Axes object.')  
assert_is_not(len(ch.title.get_text()), 0, msg="Your plot doesn't have a title.")
assert_is_not(ch.xaxis.get_label_text(), '', msg="Change the x-axis label to something more descriptive.")
assert_is_not(ch.yaxis.get_label_text(), '', msg="Change the y-axis label to something more descriptive.")
assert_equal(len(ch.patches), len(sol.patches), msg="Your bar graph doesn't have the correct number patches. Make sure you use the n_pts parameter. ")
for student_patch, solution_patch in zip(ch.patches, sol.patches):
    assert_equal(student_patch.get_width(), solution_patch.get_width(), msg='')
    assert_equal(student_patch.get_height(), solution_patch.get_height(), msg='')

```

**&copy; 2017: Robert J. Brunner at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode 